Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

### Step 1: Select and Explain the Dataset

The famous "Breast Cancer" dataset from scikit-learn. The Breast Cancer dataset contains features computed from a digitized image of a fine needle aspirate (FNA) of a breast mass. The goal is to classify tumors as malignant or benign based on various characteristics.

ชุดข้อมูลมะเร็งเต้านมประกอบด้วยคุณลักษณะที่คำนวณจากภาพดิจิทัลของเข็มดูดขนาดเล็ก (FNA) ของมวลเต้านม เป้าหมายคือการจำแนกเนื้องอกว่าเป็นเนื้อร้ายหรือไม่

### Step 2: Read Data (1 point)
Read dataset into a Pandas DataFrame.
อ่านข้อมูลจาก breast_cancer.csv เพื่อเก็บใน dataframe ชื่อ breast_cancer_data

<!-- grade_id: cell-5ce183277081aa38, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
breast_cancer_data.info()
```

<!-- grade_id: cell-0f7bf8d5311ac680, points: 1, grade: true, solution: false -->
```python
#step 2
assert len(breast_cancer_data) == 569, "Incorrect number of rows"
```

### Step 3: Exploratory Data Analysis (EDA) and Numerical Results (3 points)
คำนวณค่าสถิติพื้นฐานเพื่อสำรวจข้อมูลประกอบด้วย
* mean_radius_mean ค่าเฉลี่ยของ mean radius
* worst_perimeter_max ค่ามากสุดของ worst_perimeter
* mean_smoothness_sum ค่าผลรวมของ mean_smoothness

<!-- grade_id: cell-985e49e8cfb5ff4e, points: 0, grade: false, solution: true -->
```python
# Calculate summary statistics
# YOUR CODE HERE
raise NotImplementedError()

mean_radius_mean, worst_perimeter_max, mean_smoothness_sum
```

<!-- grade_id: cell-932f975bc5a37635, points: 3, grade: true, solution: false -->
```python
#step3
assert round(mean_radius_mean, 2) == 14.13, "Incorrect mean mean radius"
assert round(worst_perimeter_max, 2) == 251.2, "Incorrect max worst perimeter"
assert round(mean_smoothness_sum, 2) == 54.83, "Incorrect sum mean smoothness"
```

### Step 4: Setup Experiment and Model Training (3 points)

```python
from sklearn.model_selection import train_test_split
from sklearn.ensemble import RandomForestClassifier

# Split data into features (X) and target (y)
X = breast_cancer_data.drop('target', axis=1)
y = breast_cancer_data['target']
```

```python
# Split data into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
```

4.1 สร่างโมเดล RandomForestClassifier ด้วยข้อมูล X_train, y_train

<!-- grade_id: cell-0ea9d63c8557f566, points: 0, grade: false, solution: true -->
```python
# Initialize and train the random forest classifier
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c31d16161235116, points: 3, grade: true, solution: false -->
```python
#step4
# def mycheck(yp,yt,error_percent=10):
#     p=error_percent/100
#     return yt*(1-p)<yp and yp < yt*(1+p)

assert len(X_train) == 455, "Incorrect number of training samples"
assert len(X_test) == 114, "Incorrect number of test samples"

```

### Step 5: Evaluation (3 points)

<!-- grade_id: cell-6510510e34498eab, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import accuracy_score, precision_score, recall_score

# Make predictions on the test set
y_pred = model.predict(X_test)

# Calculate accuracy, precision, and recall
# YOUR CODE HERE
raise NotImplementedError()
accuracy, precision, recall
```

<!-- grade_id: cell-e13ffa942cad9259, points: 1, grade: true, solution: false -->
```python
#step5
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert mycheck(accuracy,0.965), "Incorrect accuracy"


# assert round(accuracy, 2) == 0.97, "Incorrect accuracy"
# assert round(precision, 2) == 0.98, "Incorrect precision"
# assert round(recall, 2) == 0.96, "Incorrect recall"
```

<!-- grade_id: cell-425e13b9db14985d, points: 1, grade: true, solution: false -->
```python
assert mycheck(precision,0.959), "Incorrect precision"

```

<!-- grade_id: cell-b04764de18d6ecd0, points: 1, grade: true, solution: false -->
```python
assert mycheck(recall,0.986), "Incorrect recall"
```