Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

### Step 1: Select and Explain the Dataset

California Housing Prices Dataset

ชุดข้อมูล "Boston Housing" เป็นชุดข้อมูลที่รู้จักกันดีสำหรับโจทย์ Regression ชุดข้อมูลนี้ประกอบด้วยราคาของที่อยู่อาศัยพร้อมกับฟีเจอร์ต่างๆ เช่น อัตราอาชญากรรม จำนวนห้องโดยเฉลี่ย ฯลฯ ใช้เพื่อคาดการณ์ราคาที่อยู่อาศัย

### Step 2: Read Data (1 point)
Read California Housing Prices datasetdataset into a Pandas DataFrame.

2.1 โหลด housing.csv (Boston Housing dataset) ลงใน Pandas DataFrame โดยตั้งชื่อตัวแปรของ DataFrame นั้นว่า housing_data.

<!-- grade_id: cell-5ce183277081aa38, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

housing_data.head()
```

```python
housing_data.info()
```

<!-- grade_id: cell-0f7bf8d5311ac680, points: 1, grade: true, solution: false -->
```python
#step 2.1
assert len(housing_data) == 20640, "Incorrect number of rows"
```

### Step 3: Exploratory Data Analysis (EDA) and Numerical Results (3 points)

จงคำนวณค่าทางสถิติดังต่อไปนี้
* 3.1 Mean ของ median_income
* 3.2 Max ของ housing_median_age
* 3.3 Sum ของ total_rooms

<!-- grade_id: cell-985e49e8cfb5ff4e, points: 0, grade: false, solution: true -->
```python
# Calculate summary statistics
# YOUR CODE HERE
raise NotImplementedError()

median_income_mean, housing_median_age_max, total_rooms_sum
```

<!-- grade_id: cell-932f975bc5a37635, points: 3, grade: true, solution: false -->
```python
#step3.1
assert round(median_income_mean, 2) == 3.87, "Incorrect mean median_income"
```

```python
#step3.2
assert housing_median_age_max == 52, "Incorrect max housing_median_age"
```

```python
#step3.3
assert total_rooms_sum == 54402150, "Incorrect total_rooms sum"
```

### Step 4: Setup Experiment and Model Training (3 points)

การสร้าง Regression Model ที่สามารถทำนายราคาที่อยู่อาศัย (median_house_value) สำหรับ บรรทัด ที่ 100, 200, และ 300

```python
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.impute import SimpleImputer
from sklearn.compose import ColumnTransformer
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import OneHotEncoder

# Remove the target variable
A = housing_data.drop('median_house_value', axis=1)
y = housing_data['median_house_value']
```

4.1 สร้าง matrix X ใหม่จาก A โดยแทนที่ข้อมูล ocean_proximity จากเดิมที่เป็น categorical ให้เป็น OneHotEncoding

<!-- grade_id: cell-b1eb9dc358506935, points: 0, grade: false, solution: true -->
```python
# Handle categorical variables using one-hot encoding
# YOUR CODE HERE
raise NotImplementedError()
print(X.shape)
```

<!-- grade_id: cell-6898321f55d62482, points: 2, grade: true, solution: false -->
```python
#step4.1
assert X.shape[1] == 13, "Incorrect value of shape"
```

```python
# Split data into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

```

4.2 สร้าง LinearRegression โมเดล เพื่อทำนายราคาที่อยู่อาศัย (median_house_value) สำหรับ บรรทัด ที่ 100, 200, และ 300

<!-- grade_id: cell-2e7cb46ef1f0949c, points: 0, grade: false, solution: true -->
```python
# Preprocess and train the linear regression model
# YOUR CODE HERE
raise NotImplementedError()

y_predict=model.predict(X_test[[100,200,300],:])
y_predict
```

<!-- grade_id: cell-8c31d16161235116, points: 2, grade: true, solution: false -->
```python
#step4.2
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert len(X_train) == 16512, "Incorrect number of training samples"
assert len(X_test) == 4128, "Incorrect number of test samples"
assert mycheck(y_predict[0],134173.854), "Incorrect y_prediction[0]"
assert mycheck(y_predict[1],171135.568), "Incorrect y_prediction[1]"
assert mycheck(y_predict[2],62709.449), "Incorrect y_prediction[2]"

```

### Step 5: Evaluation (3 points)

การทดลองใช้โมเดลที่ได้สร้างไว้ในขั้นตอนก่อนหน้า

5.1 ทดลองทำนายราคาที่อยู่อาศัยโดยใช้ X_test และเก็บค่าการทำนายไว้ใน y_pred

<!-- grade_id: cell-6510510e34498eab, points: 0, grade: false, solution: true -->
```python


# Make predictions on the test set
# YOUR CODE HERE
raise NotImplementedError()
print(y_pred)
```

<!-- grade_id: cell-be3a1bc24dfa91e7, points: 1, grade: true, solution: false -->
```python
#step5.1
from sklearn.metrics import mean_absolute_error, mean_squared_error, r2_score
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)
    
mse = mean_squared_error(y_test, y_pred)
assert mycheck(mse,4908290571.346)
```

5.2 คำนวณความถูกต้องของโมเดลโดยใช้มาตรวัด mean_absolute_error, mean_squared_error และ r2_score

<!-- grade_id: cell-a187fd62b97216c3, points: 0, grade: false, solution: true -->
```python
# Calculate Mean Squared Error (MSE)
# YOUR CODE HERE
raise NotImplementedError()
print(mae, mse, r2)
```

<!-- grade_id: cell-e13ffa942cad9259, points: 2, grade: true, solution: false -->
```python
#step5.2
assert  mycheck(mae,50670.489), "Incorrect Mean Absolute Error (MAE)"
assert mycheck(mse,4908290571.346), "Incorrect Mean Squared Error (MSE)"
assert mycheck(r2,0.625), "Incorrect R-squared (R2) score"
```

```python

```