### Step 1: Select and Explain the Dataset

The famous "Breast Cancer" dataset from scikit-learn. The Breast Cancer dataset contains features computed from a digitized image of a fine needle aspirate (FNA) of a breast mass. The goal is to classify tumors as malignant or benign based on various characteristics.

### Step 2: Read Data (1 point)
Read dataset into a Pandas DataFrame.

<!-- grade_id: cell-5ce183277081aa38, points: 0, grade: false, solution: true -->
```python
import pandas as pd
# from sklearn.datasets import load_breast_cancer

# Load the Breast Cancer dataset
# data = load_breast_cancer()

# breast_cancer_data = pd.DataFrame(data.data, columns=data.feature_names)
# breast_cancer_data['target'] = data.target
breast_cancer_data = pd.read_csv('breast_cancer.csv')
breast_cancer_data.head()
# breast_cancer_data.to_csv('breast_cancer.csv')
```

Output:
```
Unnamed: 0  mean radius  mean texture  mean perimeter  mean area  \
0           0        17.99         10.38          122.80     1001.0   
1           1        20.57         17.77          132.90     1326.0   
2           2        19.69         21.25          130.00     1203.0   
3           3        11.42         20.38           77.58      386.1   
4           4        20.29         14.34          135.10     1297.0   

   mean smoothness  mean compactness  mean concavity  mean concave points  \
0          0.11840           0.27760          0.3001              0.14710   
1          0.08474           0.07864          0.0869              0.07017   
2          0.10960           0.15990          0.1974              0.12790   
3          0.14250           0.28390          0.2414              0.10520   
4          0.10030           0.13280          0.1980              0.10430   

   mean symmetry  ...  worst texture  worst perimeter  worst area  \
0         0.2419  ...          17.33        
```

```python
breast_cancer_data.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 569 entries, 0 to 568
Data columns (total 32 columns):
 #   Column                   Non-Null Count  Dtype  
---  ------                   --------------  -----  
 0   Unnamed: 0               569 non-null    int64  
 1   mean radius              569 non-null    float64
 2   mean texture             569 non-null    float64
 3   mean perimeter           569 non-null    float64
 4   mean area                569 non-null    float64
 5   mean smoothness          569 non-null    float64
 6   mean compactness         569 non-null    float64
 7   mean concavity           569 non-null    float64
 8   mean concave points      569 non-null    float64
 9   mean symmetry            569 non-null    float64
 10  mean fractal dimension   569 non-null    float64
 11  radius error             569 non-null    float64
 12  texture error            569 non-null    float64
 13  perimeter error          569 non-null    float64
 14  area error               5
```

<!-- grade_id: cell-0f7bf8d5311ac680, points: 1, grade: true, solution: false -->
```python
#step 2
assert len(breast_cancer_data) == 569, "Incorrect number of rows"
```

### Step 3: Exploratory Data Analysis (EDA) and Numerical Results (3 points)

<!-- grade_id: cell-985e49e8cfb5ff4e, points: 0, grade: false, solution: true -->
```python
# Calculate summary statistics
mean_radius_mean = breast_cancer_data['mean radius'].mean()
worst_perimeter_max = breast_cancer_data['worst perimeter'].max()
mean_smoothness_sum = breast_cancer_data['mean smoothness'].sum()

mean_radius_mean, worst_perimeter_max, mean_smoothness_sum
```

Output:
```
(14.127291739894552, 251.2, 54.82900000000001)
```

<!-- grade_id: cell-932f975bc5a37635, points: 3, grade: true, solution: false -->
```python
#step3
assert round(mean_radius_mean, 2) == 14.13, "Incorrect mean mean radius"
assert round(worst_perimeter_max, 2) == 251.2, "Incorrect max worst perimeter"
assert round(mean_smoothness_sum, 2) == 54.83, "Incorrect sum mean smoothness"
```

### Step 4: Setup Experiment and Model Training (3 points)

<!-- grade_id: cell-c7ace6ade31e8be3, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
from sklearn.ensemble import RandomForestClassifier

# Split data into features (X) and target (y)
X = breast_cancer_data.drop('target', axis=1)
y = breast_cancer_data['target']

# Split data into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

# Initialize and train the random forest classifier
model = RandomForestClassifier(random_state=42)
model.fit(X_train, y_train)
```

Output:
```
RandomForestClassifier(random_state=42)
```

<!-- grade_id: cell-8c31d16161235116, points: 3, grade: true, solution: false -->
```python
#step4
# def mycheck(yp,yt,error_percent=10):
#     p=error_percent/100
#     return yt*(1-p)<yp and yp < yt*(1+p)

assert len(X_train) == 455, "Incorrect number of training samples"
assert len(X_test) == 114, "Incorrect number of test samples"

```

### Step 5: Evaluation (3 points)

<!-- grade_id: cell-6510510e34498eab, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import accuracy_score, precision_score, recall_score

# Make predictions on the test set
y_pred = model.predict(X_test)

# Calculate accuracy, precision, and recall
accuracy = accuracy_score(y_test, y_pred)
precision = precision_score(y_test, y_pred)
recall = recall_score(y_test, y_pred)

accuracy, precision, recall
```

Output:
```
(0.9649122807017544, 0.958904109589041, 0.9859154929577465)
```

<!-- grade_id: cell-e13ffa942cad9259, points: 3, grade: true, solution: false -->
```python
#step5
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert mycheck(accuracy,0.965), "Incorrect accuracy"
assert mycheck(precision,0.959), "Incorrect precision"
assert mycheck(recall,0.986), "Incorrect recall"

# assert round(accuracy, 2) == 0.97, "Incorrect accuracy"
# assert round(precision, 2) == 0.98, "Incorrect precision"
# assert round(recall, 2) == 0.96, "Incorrect recall"
```

```python

```