Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Wasit"
COLLABORATORS = ""
```

---

### Step 1: Select and Explain the Dataset

The famous "Breast Cancer" dataset from scikit-learn. The Breast Cancer dataset contains features computed from a digitized image of a fine needle aspirate (FNA) of a breast mass. The goal is to classify tumors as malignant or benign based on various characteristics.

### Step 2: Read Data (1 point)
Read dataset into a Pandas DataFrame.

<!-- grade_id: cell-5ce183277081aa38, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
breast_cancer_data = pd.read_csv('breast_cancer.csv')
breast_cancer_data.head()
```

```python
breast_cancer_data.info()
```

<!-- grade_id: cell-0f7bf8d5311ac680, points: 1, grade: true, solution: false -->
```python
#step 2
assert len(breast_cancer_data) == 569, "Incorrect number of rows"
```

### Step 3: Exploratory Data Analysis (EDA) and Numerical Results (3 points)

<!-- grade_id: cell-985e49e8cfb5ff4e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-932f975bc5a37635, points: 3, grade: true, solution: false -->
```python
#step3
assert round(mean_radius_mean, 2) == 14.13, "Incorrect mean mean radius"
assert round(worst_perimeter_max, 2) == 251.2, "Incorrect max worst perimeter"
assert round(mean_smoothness_sum, 2) == 54.83, "Incorrect sum mean smoothness"
```

### Step 4: Setup Experiment and Model Training (3 points)

<!-- grade_id: cell-c7ace6ade31e8be3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8c31d16161235116, points: 3, grade: true, solution: false -->
```python
#step4
# def mycheck(yp,yt,error_percent=10):
#     p=error_percent/100
#     return yt*(1-p)<yp and yp < yt*(1+p)

assert len(X_train) == 455, "Incorrect number of training samples"
assert len(X_test) == 114, "Incorrect number of test samples"

```

### Step 5: Evaluation (3 points)

<!-- grade_id: cell-6510510e34498eab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e13ffa942cad9259, points: 3, grade: true, solution: false -->
```python
#step5
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert mycheck(accuracy,0.965), "Incorrect accuracy"
assert mycheck(precision,0.959), "Incorrect precision"
assert mycheck(recall,0.986), "Incorrect recall"

# assert round(accuracy, 2) == 0.97, "Incorrect accuracy"
# assert round(precision, 2) == 0.98, "Incorrect precision"
# assert round(recall, 2) == 0.96, "Incorrect recall"
```

```python

```