### Step 1: The Dataset

ชุดข้อมูล "Boston Housing" เป็นชุดข้อมูลที่รู้จักกันดีสำหรับโจทย์ Regression ชุดข้อมูลนี้ประกอบด้วยราคาของที่อยู่อาศัยพร้อมกับฟีเจอร์ต่างๆ เช่น อัตราอาชญากรรม จำนวนห้องโดยเฉลี่ย ฯลฯ ใช้เพื่อคาดการณ์ราคาที่อยู่อาศัย

### Step 2: Read Data (1 point)
โหลด housing.csv (Boston Housing dataset) ลงใน Pandas DataFrame โดยตั้งชื่อตัวแปรของ DataFrame นั้นว่า housing_data.

<!-- grade_id: cell-5ce183277081aa38, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
import pandas as pd
housing_data = pd.read_csv('housing.csv')
### END SOLUTION
print(housing_data.head())
print(housing_data.info())
```

Output:
```
longitude  latitude  housing_median_age  total_rooms  total_bedrooms  \
0    -122.23     37.88                41.0        880.0           129.0   
1    -122.22     37.86                21.0       7099.0          1106.0   
2    -122.24     37.85                52.0       1467.0           190.0   
3    -122.25     37.85                52.0       1274.0           235.0   
4    -122.25     37.85                52.0       1627.0           280.0   

   population  households  median_income  median_house_value ocean_proximity  
0       322.0       126.0         8.3252            452600.0        NEAR BAY  
1      2401.0      1138.0         8.3014            358500.0        NEAR BAY  
2       496.0       177.0         7.2574            352100.0        NEAR BAY  
3       558.0       219.0         5.6431            341300.0        NEAR BAY  
4       565.0       259.0         3.8462            342200.0        NEAR BAY  
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 20640 entries, 0 to 20639
Da
```

<!-- grade_id: cell-01e8ae97151665e4, points: 1, grade: true, solution: false -->
```python
assert len(housing_data) == 20640, "Incorrect number of rows"
```

### Step 3: Exploratory Data Analysis (EDA) and Numerical Results (3 points)
จงคำนวณค่าทางสถิติดังต่อไปนี้
- Mean ของ median_income
- Max ของ housing_median_age
- Sum ของ total_rooms

<!-- grade_id: cell-985e49e8cfb5ff4e, points: 0, grade: false, solution: true -->
```python
# Calculate summary statistics

### BEGIN SOLUTION
median_income_mean = housing_data['median_income'].mean()
housing_median_age_max = housing_data['housing_median_age'].max()
total_rooms_sum = housing_data['total_rooms'].sum()
### END SOLUTION

print(median_income_mean)
print(housing_median_age_max)
print(total_rooms_sum)
```

Output:
```
3.8706710029069766
52.0
54402150.0
```

<!-- grade_id: cell-932f975bc5a37635, points: 3, grade: true, solution: false -->
```python
assert round(median_income_mean, 2) == 3.87, "Incorrect mean median_income"
assert housing_median_age_max == 52, "Incorrect max housing_median_age"
assert total_rooms_sum == 54402150, "Incorrect total_rooms sum"
```

### Step 4: Setup Experiment and Model Training (3 points)
สร้าง Regression Model ที่สามารถทำนายราคาที่อยู่อาศัย (median_house_value) สำหรับ บรรทัด ที่ 100, 200, และ 300

```python
# Import important libraries

from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
from sklearn.impute import SimpleImputer
from sklearn.compose import ColumnTransformer
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import OneHotEncoder

# Remove the target variable
A = housing_data.drop('median_house_value', axis=1)
y = housing_data['median_house_value']
```

### สร้าง matrix X ใหม่จาก A โดยแทนที่ข้อมูล ocean_proximity จากเดิมที่เป็น categorical ให้เป็น OneHotEncoding

```python
### BEGIN SOLUTION
categorical_features = ['ocean_proximity']
numeric_features = list(A.columns.difference(categorical_features))
preprocessor = ColumnTransformer(
    transformers=[
        ('num', SimpleImputer(strategy='median'), numeric_features),
        ('cat', OneHotEncoder(), categorical_features)
    ])

X = preprocessor.fit_transform(A)
### END SOLUTION
print(X.shape)
```

Output:
```
(20640, 13)
```

<!-- grade_id: cell-4eec39ae19a2abfe, points: 2, grade: true, solution: false -->
```python
assert X.shape[1] == 13, "Incorrect value of shape"
```

<!-- grade_id: cell-c7ace6ade31e8be3, points: 0, grade: false, solution: true -->
```python
# Split data into training and testing sets
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)

# Preprocess and train the linear regression model
### BEGIN SOLUTION
model = Pipeline(steps=[
    ('preprocessor', preprocessor),
    ('model', LinearRegression())
])
### END SOLUTION

model.fit(X_train, y_train)
y_predict=model.predict(X_test.iloc[[100,200,300],:])
y_predict
```

Output:
```
File ~/anaconda3/lib/python3.11/site-packages/sklearn/compose/_column_transformer.py:448, in ColumnTransformer._validate_column_callables(self, X)
    446         columns = columns(X)
    447     all_columns.append(columns)
--> 448     transformer_to_input_indices[name] = _get_column_indices(X, columns)
    450 self._columns = all_columns
    451 self._transformer_to_input_indices = transformer_to_input_indices

File ~/anaconda3/lib/python3.11/site-packages/sklearn/utils/__init__.py:425, in _get_column_indices(X, key)
    423     all_columns = X.columns
    424 except AttributeError:
--> 425     raise ValueError(
    426         "Specifying the columns using strings is only "
    427         "supported for pandas DataFrames"
    428     )
    429 if isinstance(key, str):
    430     columns = [key]

ValueError: Specifying the columns using strings is only supported for pandas DataFrames
```

<!-- grade_id: cell-21b208fdf24a7ad1, points: 1, grade: true, solution: false -->
```python
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert len(X_train) == 16512, "Incorrect number of training samples"
assert len(X_test) == 4128, "Incorrect number of test samples"
assert mycheck(y_predict[0],134173.854), "Incorrect y_prediction[0]"
assert mycheck(y_predict[1],171135.568), "Incorrect y_prediction[1]"
assert mycheck(y_predict[2],215393.475), "Incorrect y_prediction[2]"

```

### Step 5: Evaluation (3 points)

<!-- grade_id: cell-6510510e34498eab, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_absolute_error, mean_squared_error, r2_score

# Make predictions on the test set
y_pred = model.predict(X_test)

# Calculate Mean Squared Error (MSE)
### BEGIN SOLUTION
mae = mean_absolute_error(y_test, y_pred)
mse = mean_squared_error(y_test, y_pred)
r2 = r2_score(y_test, y_pred)
### END SOLUTION
mae, mse, r2
```

Output:
```
(50670.48923565551, 4908290571.346386, 0.6254382675296302)
```

<!-- grade_id: cell-e13ffa942cad9259, points: 3, grade: true, solution: false -->
```python
#step5
def mycheck(yp,yt,error_percent=10):
    p=error_percent/100
    return yt*(1-p)<yp and yp < yt*(1+p)

assert  mycheck(mae,50670.489), "Incorrect Mean Absolute Error (MAE)"
assert mycheck(mse,4908290571.346), "Incorrect Mean Squared Error (MSE)"
assert mycheck(r2,0.625), "Incorrect R-squared (R2) score"
```